# Graded Assessment

In this assessment you will write a full end-to-end training process using gluon and mxnet. We will train the LeNet-5 classifier network on the MNIST dataset. The network will be defined for you but you have to fill in code to prepare the dataset, train the network, and evaluate it's performance on a held out dataset.

```python
from pathlib import Path
from mxnet import gluon, metric, autograd, init, nd
import os
```

```python
M5_DATA = Path(os.getenv('DATA_DIR', '../../data'), 'module_5')
M5_IMAGES = Path(M5_DATA, 'images')
```

---
## Question 1

### Prepare and the data and construct the dataloader

* First, get the MNIST dataset from `gluon.data.vision.datasets`. Use
* Don't forget the ToTensor and normalize Transformations. Use `0.13` and `0.31` as the mean and standard deviation respectively
* Construct the dataloader with the batch size provide. Ensure that the train_dataloader is shuffled.

Note: Please specify `root=M5_IMAGES` for your dataset storage location.

<!-- grade_id: dataloader, points: 0, grade: false, solution: true -->
```python
import os
from pathlib import Path

def get_mnist_data(batch=128):
    """
    Should construct a dataloader with the MNIST Dataset with the necessary transforms applied.
    
    :param batch: batch size for the DataLoader.
    :type batch: int
    
    :return: a tuple of the training and validation DataLoaders
    :rtype: (gluon.data.DataLoader, gluon.data.DataLoader)
    """
    train_data = gluon.data.vision.datasets.MNIST(root=M5_IMAGES, train=True)
    val_data = gluon.data.vision.datasets.MNIST(root=M5_IMAGES, train=False)
    
    transforms = gluon.data.vision.transforms
    
    transform_fn = transforms.Compose([transforms.ToTensor(),
                                  transforms.Normalize(0.13, 0.31)])
    
    train_dataloader = gluon.data.DataLoader(dataset=train_data.transform_first(transform_fn),
                                             batch_size=batch,
                                            shuffle=True)
    
    validation_dataloader = gluon.data.DataLoader(dataset=val_data.transform_first(transform_fn),
                                             batch_size=batch,
                                            shuffle=False)
    
    return train_dataloader, validation_dataloader
```

<!-- grade_id: cell-3ec320323b376193, points: 3, grade: true, solution: false -->
```python
t, v = get_mnist_data()
assert isinstance(t, gluon.data.DataLoader)
assert isinstance(v, gluon.data.DataLoader)

d, l = next(iter(t))
assert d.shape == (128, 1, 28, 28) #check Channel First and Batch Size
assert l.shape == (128,)

assert nd.max(d).asscalar() <= 2.9 # check for normalization
assert nd.min(d).asscalar() >= -0.5 # check for normalization
```

---

## Question 2

### Write the training loop

* Create the loss function. This should be a loss function suitable for multi-class classification.
* Create the metric accumulator. This should the compute and store the accuracy of the model during training
* Create the trainer with the `adam` optimizer and learning rate of `0.002`
* Write the training loop

<!-- grade_id: train, points: 0, grade: false, solution: true -->
```python
def train(network, training_dataloader, batch_size, epochs):
    """
    Should take an initialized network and train that network using data from the data loader.
    
    :param network: initialized gluon network to be trained
    :type network: gluon.Block
    
    :param training_dataloader: the training DataLoader provides batches for data for every iteration
    :type training_dataloader: gluon.data.DataLoader
    
    :param batch_size: batch size for the DataLoader.
    :type batch_size: int
    
    :param epochs: number of epochs to train the DataLoader
    :type epochs: int
    
    :return: tuple of trained network and the final training accuracy
    :rtype: (gluon.Block, float)
    """
    
    softmax = gluon.loss.SoftmaxCrossEntropyLoss()
    training_accuracy = metric.Accuracy()
    
    trainer = gluon.Trainer(network.collect_params(),
                            optimizer='adam', optimizer_params={'learning_rate':0.002})
    
    for epoch in range(epochs):
        train_loss = 0.0

        for data, label in training_dataloader:
            with autograd.record():
                output = network(data)
                loss = softmax(output, label)
                
            loss.backward()
            
            trainer.step(batch_size)
            
            train_loss += loss.mean().asscalar()
            training_accuracy.update(label, output)
            
        print('Epoch[%d] Loss: %.3f Acc %.3f',
             epoch,
              train_loss/len(training_dataloader),
              training_accuracy.get()[1])

    return network, training_accuracy.get()[1]
```

Let's define and initialize a network to test the train function

```python
net = gluon.nn.Sequential()
net.add(gluon.nn.Conv2D(channels=6, kernel_size=5, activation='relu'),
        gluon.nn.MaxPool2D(pool_size=2, strides=2),
        gluon.nn.Conv2D(channels=16, kernel_size=3, activation='relu'),
        gluon.nn.MaxPool2D(pool_size=2, strides=2),
        gluon.nn.Flatten(),
        gluon.nn.Dense(120, activation="relu"),
        gluon.nn.Dense(84, activation="relu"),
        gluon.nn.Dense(10))
net.initialize(init=init.Xavier())
```

<!-- grade_id: cell-9d116b0901b24fcc, points: 5, grade: true, solution: false -->
```python
n, ta = train(net, t, 128, 5)
assert ta >= .95

d, l = next(iter(v))
p = (n(d).argmax(axis=1))
assert (p.asnumpy() == l.asnumpy()).sum()/128.0 > .95
```

Output:
```
Epoch[%d] Loss: %.3f Acc %.3f 0 0.21463440837406084 0.9332666666666667
Epoch[%d] Loss: %.3f Acc %.3f 1 0.06485772433418677 0.9564666666666667
Epoch[%d] Loss: %.3f Acc %.3f 2 0.0482370805641664 0.9658888888888889
Epoch[%d] Loss: %.3f Acc %.3f 3 0.03865339931672507 0.9714625
Epoch[%d] Loss: %.3f Acc %.3f 4 0.03134737699069202 0.97514
TypeError                                 Traceback (most recent call last)
<ipython-input-20-5e29885caea6> in <module>
      1 n, ta = train(net, t, 128, 5)
----> 2 assert ta >= .95
      3 
      4 d, l = next(iter(v))
      5 p = (n(d).argmax(axis=1))

TypeError: '>=' not supported between instances of 'Accuracy' and 'float'
```

---
## Question 3

### Write the validation loop

* Create the metric accumulator. This should the compute and store the accuracy of the model on the validation set
* Write the validation loop

<!-- grade_id: sum_of_squares_application, points: 0, grade: false, solution: true -->
```python
def validate(network, validation_dataloader):
    """
    Should compute the accuracy of the network on the validation set.
    
    :param network: initialized gluon network to be trained
    :type network: gluon.Block
    
    :param validation_dataloader: the training DataLoader provides batches for data for every iteration
    :type validation_dataloader: gluon.data.DataLoader
    
    :return: validation accuracy
    :rtype: float
    """
    validation_accuracy = metric.Accuracy()
    
    for data, label in validation_dataloader:
        output = network(data)
        validation_accuracy.update(label, output)
        
    print('Validation accuracy %.3f', validation_accuracy.get()[1]) 
    
    return validation_accuracy.get()[1]
```

<!-- grade_id: cell-5a2130f381810258, points: 2, grade: true, solution: false -->
```python
assert validate(n, v) > .95
```

Output:
```
Validation accuracy %.3f 0.9826
```

```python

```